---
title: 'Por que Python?'
description: '"Professor, qual a melhor linguagem de programação?"...Existe resposta para tal questionamento? Por que mais uma linguagem de programação? Por que Python?'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: c4ac60420e
xp: 50
```

`@projector_key`
93f7876476f6a238563ee0b4ffd83fcf
